## Pandas Lab

In this assignment, we'll be looking for outliers in petal and sepal length amongst the iris dataset. The iris dataset is a classic in machine learning and has petal and sepal data for different iris species. Let's get started.

```python
# importing modules

from sklearn import datasets
import pandas as pd
import numpy as np
```

```python
# We can load in the iris dataset directly from sklearn.

iris = datasets.load_iris()
```

```python
# The petal and sepal data is stored in "iris.data". Note that it's in a 2-D array.
# The columns correspond to ['sepal_length', 'sepal_width', 'petal_length', 'petal_width']

iris.data
```

Output:
```
array([[5.1, 3.5, 1.4, 0.2],
       [4.9, 3. , 1.4, 0.2],
       [4.7, 3.2, 1.3, 0.2],
       [4.6, 3.1, 1.5, 0.2],
       [5. , 3.6, 1.4, 0.2],
       [5.4, 3.9, 1.7, 0.4],
       [4.6, 3.4, 1.4, 0.3],
       [5. , 3.4, 1.5, 0.2],
       [4.4, 2.9, 1.4, 0.2],
       [4.9, 3.1, 1.5, 0.1],
       [5.4, 3.7, 1.5, 0.2],
       [4.8, 3.4, 1.6, 0.2],
       [4.8, 3. , 1.4, 0.1],
       [4.3, 3. , 1.1, 0.1],
       [5.8, 4. , 1.2, 0.2],
       [5.7, 4.4, 1.5, 0.4],
       [5.4, 3.9, 1.3, 0.4],
       [5.1, 3.5, 1.4, 0.3],
       [5.7, 3.8, 1.7, 0.3],
       [5.1, 3.8, 1.5, 0.3],
       [5.4, 3.4, 1.7, 0.2],
       [5.1, 3.7, 1.5, 0.4],
       [4.6, 3.6, 1. , 0.2],
       [5.1, 3.3, 1.7, 0.5],
       [4.8, 3.4, 1.9, 0.2],
       [5. , 3. , 1.6, 0.2],
       [5. , 3.4, 1.6, 0.4],
       [5.2, 3.5, 1.5, 0.2],
       [5.2, 3.4, 1.4, 0.2],
       [4.7, 3.2, 1.6, 0.2],
       [4.8, 3.1, 1.6, 0.2],
       [5.4, 3.4, 1.5, 0.4],
       [5.2, 4.1, 1.5, 0.1],
       [5.5, 4.2, 1.4, 0.2],
       [4.9, 3
```

```python
# The species data is stored in "iris.target". 

iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

## Part 1)
a) 10 pts

b) 20 pts

<!-- grade_id: cell-52d07369c60897fe, points: 0, grade: false, solution: true -->
```python
'''
a) Create a pandas dataframe called 'data' that combines the petal and sepal data with species information 
The five columns should be 'sepal_length', 'sepal_width', 'petal_length', 'petal_width', and 'species'
'''

# your code here
data = pd.DataFrame(iris.data, columns = ['sepal_length','sepal_width','petal_length','petal_width'])
target = pd.DataFrame(iris.target,columns = ['species'])
data = pd.concat([data, target], axis=1)
```

<!-- grade_id: cell-5fa1c6760c60a557, points: 10, grade: true, solution: false -->
```python
'''
END OF SECTION
'''
```

<!-- grade_id: cell-ec6e7aa55d648407, points: 0, grade: false, solution: true -->
```python
'''
b) Notice how the 'species' column is just numbers. The names of the species are stored in a list called target_names.
Replace the species number with the corresponding name from the target_names list.
HINT: We can use the df.replace() function with a map dictionary.
'''

# your code here
species_map = [  data['species'] == 0,
                 data['species'] == 1,
                 data['species'] == 2 ]
target_names = ['setosa','versicolor','virginica']
data['species'] = np.select(species_map,target_names,'oth')
```

```python
data['species'].value_counts()
```

Output:
```
virginica     50
setosa        50
versicolor    50
Name: species, dtype: int64
```

<!-- grade_id: cell-3c0a74ac20689049, points: 20, grade: true, solution: false -->
```python
'''
END OF SECTION
'''

### BEGIN HIDDEN TEST
assert data['species'][10] == 'setosa'
assert data['species'][100] == 'virginica'
assert data['species'][75] == 'versicolor'
### END HIDDEN TEST
```

## PART 2
a) 5 pts

b) 25 pts

c) 10 pts

Our goal is to find outliers in the data. We want to know which rows have length/width ratios that deviate far from the avg. value of its species.  

<!-- grade_id: cell-45df183403b3ce40, points: 0, grade: false, solution: true -->
```python
'''
a) First, we have to create two new columns, 'sepal_length/width'
and 'petal length/width' to get a sense of the ratios. Create these columns in our 'data' dataframe.
'''

# your code here
data['sepal_length/width'] = data['sepal_length'] / data['sepal_width']
data['petal_length/width'] = data['petal_length'] / data['petal_width']
```

<!-- grade_id: cell-5566d05e6e24d518, points: 5, grade: true, solution: false -->
```python
'''
END OF SECTION
'''

### BEGIN HIDDEN TEST
assert data['sepal_length/width'][5] == 1.3846153846153848
assert data['petal_length/width'][5] == 4.25
### END HIDDEN TEST
```

<!-- grade_id: cell-2ac05314e7bc6cb2, points: 0, grade: false, solution: true -->
```python
'''
b) Now, use the transform function to find the 'mean_sepal_length/width'
and the 'mean_petal_length/width' for each species. You should create these columns
in the 'data' dataframe.
'''

# your code here
mn_sepal = data.groupby(['species'])['sepal_length/width'].mean() 
mn_petal = data.groupby(['species'])['petal_length/width'].mean()
print(mn_sepal)
print(mn_petal)
```

Output:
```
species
setosa        1.470188
versicolor    2.160402
virginica     2.230453
Name: sepal_length/width, dtype: float64
species
setosa        6.908000
versicolor    3.242837
virginica     2.780662
Name: petal_length/width, dtype: float64
```

```python
mn_sepal = pd.DataFrame(mn_sepal)
mn_petal = pd.DataFrame(mn_petal)
```

```python
mn_sepal
```

Output:
```
sepal_length/width
species                       
setosa                1.470188
versicolor            2.160402
virginica             2.230453
```

```python
data.head()
```

Output:
```
sepal_length  sepal_width  petal_length  petal_width species  \
0           5.1          3.5           1.4          0.2  setosa   
1           4.9          3.0           1.4          0.2  setosa   
2           4.7          3.2           1.3          0.2  setosa   
3           4.6          3.1           1.5          0.2  setosa   
4           5.0          3.6           1.4          0.2  setosa   

   sepal_length/width  petal_length/width  
0            1.457143                 7.0  
1            1.633333                 7.0  
2            1.468750                 6.5  
3            1.483871                 7.5  
4            1.388889                 7.0
```

```python
data = data.join(mn_sepal, on='species', how='left', rsuffix='_mean')
```

```python
data = data.join(mn_petal, on='species', how='left', rsuffix='_mean')
```

```python
data = data.rename(columns={"sepal_length/width_mean": "mean_sepal_length/width", "petal_length/width_mean": "mean_petal_length/width"})
```

<!-- grade_id: cell-c1bff42506d647aa, points: 25, grade: true, solution: false -->
```python
'''
END OF SECTION
'''

### BEGIN HIDDEN TEST
assert data['mean_sepal_length/width'][5] == 1.4701876810227483
assert data['mean_petal_length/width'][5] == 6.9079999999999995
### END HIDDEN TEST
```

<!-- grade_id: cell-69006321e7d1527a, points: 0, grade: false, solution: true -->
```python
'''
c) Find the squared difference between each row's length/width value and the average of its species.
Create two new columns titled 'sq_diff_sepal_length/width' and ''sq_diff_petal_length/width'. 
HINT: Use the function np.square()
'''

# your code here
data['sq_diff_sepal_length/width'] = np.square(data['sepal_length/width'] - data['mean_sepal_length/width'])
data['sq_diff_petal_length/width'] = np.square(data['petal_length/width'] - data['mean_petal_length/width'])
```

<!-- grade_id: cell-992c2ec5955fecd1, points: 10, grade: true, solution: false -->
```python
'''
END OF SECTION
'''

### BEGIN HIDDEN TEST
assert data['sq_diff_sepal_length/width'][5] == 0.0073226179124296804
assert data['sq_diff_petal_length/width'][5] == 7.064963999999997
### END HIDDEN TEST
```

## PART 3
a) 15 pts

b) 15 pts

<!-- grade_id: cell-27a18a1f95e90753, points: 0, grade: false, solution: true -->
```python
'''
a) Create a new dataframe called 'outliers_sepal' and 'outliers_petal' with the top 20 values that deviate most.
HINT: Use the sort_values() function.
'''

# your code here
outliers_sepal = data.sort_values(by = 'sq_diff_sepal_length/width',ascending=False).head(20)
outliers_petal = data.sort_values(by = 'sq_diff_petal_length/width',ascending=False).head(20)
```

<!-- grade_id: cell-e41ba6b19ca18e85, points: 15, grade: true, solution: false -->
```python
'''
END OF SECTION
'''

### BEGIN HIDDEN TEST
assert outliers_sepal.iloc[0, 1] == 2.6
assert outliers_petal.iloc[0, 1] == 4.1
### END HIDDEN TEST
```

<!-- grade_id: cell-ed873617fefbab4e, points: 0, grade: false, solution: true -->
```python
'''
b) Do any of the rows overlap? Create a list called 'outliers' to store the index numbers of any overlapping rows.
HINT: Use list comprehension.
'''

# your code here
outliers = [x for x in outliers_petal.index.values if x in outliers_sepal.index.values]
```

<!-- grade_id: cell-070c4c9d4d32f89d, points: 15, grade: true, solution: false -->
```python
'''
END OF SECTION
'''

### BEGIN HIDDEN TEST
assert outliers == [41]
### END HIDDEN TEST
```